```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import seaborn as sns
import pandas as pd
import numpy as np
import numpy.testing as nt
```

# Optimisation II - Lab 2


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.14.5/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!

Developed by Dr M. Woolway and edited by Krupa Prag. 

### Question 2 (1 Mark)

Write a function  that takes as input two positive integer scalars, $\texttt{n}$ and $\texttt{m}$, in that order. The function must create and return board, which is an n-by-m matrix. Every element of board is either 0 or 1. The first element, board(0,0) is 1. No direct neighbours in the matrix, vertically or horizontally, can be equal. That is, a 1 element cannot have 1 immediately preceding or following it in the same row or column.

<!-- grade_id: cell-a6ee6e1a52be2a1f, points: 0, grade: false, solution: true -->
```python
def check_board(n, m):
    """
    Inputs:
        n, m: scalar integers
    Outputs:
        board: binary numpy matrix
    """
    # YOUR CODE HERE
    
    array = np.zeros(n*m, dtype=int)
    array[::2]=1
    
    return array.reshape(n,m)

    raise NotImplementedError()
```

<!-- grade_id: cell-758afda387a26ce1, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
test1 = np.array([[1., 0., 1., 0., 1.],
                  [0., 1., 0., 1., 0.],
                  [1., 0., 1., 0., 1.],
                  [0., 1., 0., 1., 0.]])
nt.assert_array_almost_equal(test1, check_board(4, 5), err_msg="Function Incorrect")
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

### Question 3 (1 Mark)

Write a function that takes three input arguments: $\texttt{limit}$, $\texttt{n}$, and $\texttt{m}$, in that order. The function returns an n-by-m matrix of uniformly distributed random integers between 1 and limit inclusive.  You are not allowed to use randi, but you can use rand. You will also find the built-in function floor useful for obtaining integers. To make sure that your result is indeed uniformly distributed, test the output of the function by using the built-in function hist, which plots a histogram on larger inputs.

<!-- grade_id: cell-da3241867a22aab7, points: 0, grade: false, solution: true -->
```python
import random as random
def random_dist(limit, n, m):
    # YOUR CODE HERE
    return [random.choices(range(1,limit), k=m)for _ in range(n)]

    raise NotImplementedError()
```

<!-- grade_id: cell-cdf31d49bdc2ef61, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
test1 = random_dist(5, 5, 5)
def isinteger(x):
     return np.equal(np.mod(x, 1), 0)
assert(np.all(isinteger(test1)) == True)
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

# Main Exercises

### Question 1 (3 Marks)

Write a function that implements the Secant method. It should take as inputs the function handle $\texttt{f}$ (derivative function), two initial guesses $x0, \ x1$ and a tolerance $\texttt{tol}$. The function should return the root of the function. Use the absolute difference to control your tolerance.
 

<!-- grade_id: cell-be8be7d876524594, points: 0, grade: false, solution: true -->
```python
def secant_min(f, x0, x1, tol):
    
    """
    Inputs:
        f : a function handle
        x0: an initial point - scalar
        x1: a subsequent point - scalar
        
    Outputs:
        x: root of function f - scalar
    """
    while(abs(x0-x1)>tol):
        x1_prev=x1
        x1=x0-f(x0)*(x1-x0)/(f(x1)-f(x0))
        x0=x1_prev
    
    return x0
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c10a56da7c9af328, points: 3, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
f   = lambda x: x**2 - 612
x0  = 10
x1  = 30
tol = 10**(-5)
nt.assert_almost_equal(24.738633753705965, secant_min(f, x0, x1, tol), err_msg='incorrect function')
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

### Question 2 (4 Marks)

Write a function that performs the Golden Search method. It should take in the function handle $\texttt{f}$, the domain, $\texttt{a}$, $\texttt{b}$ and a tolerance $\texttt{tol}$. It should return a both the function value at the minima and the minima itself. Use the absolute difference to control your tolerance.

<!-- grade_id: cell-f7e3789958524bb4, points: 0, grade: false, solution: true -->
```python
def golden_search(f, a, b, tol):
    """
    Inputs:
        f : a function handle
        a : left endpoint - scalar
        b : right endpoint - scalar
        
    Outputs:
        out: np.array of xmin and f(xmin)
    """
    # YOUR CODE HERE
    goldenRatio = (sqrt(5) + 1) / 2
    a_1 = b - (b - a) / goldenRatio
    b_1 = a + (b - a) / goldenRatio
    
    while abs(b - a) > tol:
        if f(a_1) > f(b_1):
            a = a_1
        else:
            b = b_1
        
        b_1 = a + (b - a) / goldenRatio
        a_1 = b - (b - a) / goldenRatio
        
 
    minVal = (b + a) / 2
    minn = f(minVal)
    return np.append(minVal,minn)
    raise NotImplementedError()
```

<!-- grade_id: cell-9a107226a9b40cdb, points: 4, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
f = lambda x: x**4 - 14*x**3 + 60*x**2 - 70*x
a = 0
b = 2
tol = 10**(-5)
nt.assert_array_almost_equal(np.array([0.78088186, -24.36960157]), golden_search(f, a, b, tol), err_msg='incorrect function')
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

```python

```