```python
# Initialisation Cell
from matplotlib import pyplot as plt
from IPython.display import display, HTML, Javascript
from math import *
import seaborn as sns
import pandas as pd
import numpy as np
import numpy.testing as nt
import time
```

# Optimisation II - Lab 3


## Instructions

* Read all the instructions carefully.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.14.5/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!
* To launch jupyter correctly in the labs using vmuser, type the following at the terminal:
    `/usr/local/anaconda3/bin/jupyter-notebook`
    
Developed by Dr M. Woolway and edited by Krupa Prag.     

## Warm-up Exercises



### Question 1

Write a function that called performs a multiplication table. The function takes in inputs, $n$ and $m$. The function outputs a $n\times m$  matrix which is the resultant multiplication table. That is, the element at row $i$ and column $j$ equals to $i\times j$. The function is also to have a second output which is the sum of all elements of the matrix.


<!-- grade_id: cell-4394ee0f7b563da6, points: 0, grade: false, solution: true -->
```python
def multi_table(n, m):
    # YOUR CODE HERE
    size=(n,m)
    arr=np.zeros(size)
    sum_t=0
    
    for i in range(n):
        for j in range(m):
            arr[i][j]=(i+1)*(j+1)
            sum_t+=(i+1)*(j+1)
            
    return arr, sum_t
    
    raise NotImplementedError()
```

<!-- grade_id: cell-1642bed4110ddd77, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Mark
o1 = np.array([[  1.,   2.,   3.,   4.,   5.,   6.,   7.,   8.,   9.,  10.],
               [  2.,   4.,   6.,   8.,  10.,  12.,  14.,  16.,  18.,  20.],
               [  3.,   6.,   9.,  12.,  15.,  18.,  21.,  24.,  27.,  30.],
               [  4.,   8.,  12.,  16.,  20.,  24.,  28.,  32.,  36.,  40.],
               [  5.,  10.,  15.,  20.,  25.,  30.,  35.,  40.,  45.,  50.],
               [  6.,  12.,  18.,  24.,  30.,  36.,  42.,  48.,  54.,  60.],
               [  7.,  14.,  21.,  28.,  35.,  42.,  49.,  56.,  63.,  70.],
               [  8.,  16.,  24.,  32.,  40.,  48.,  56.,  64.,  72.,  80.],
               [  9.,  18.,  27.,  36.,  45.,  54.,  63.,  72.,  81.,  90.],
               [ 10.,  20.,  30.,  40.,  50.,  60.,  70.,  80.,  90., 100.],
               [ 11.,  22.,  33.,  44.,  55.,  66.,  77.,  88.,  99., 110.],
               [ 12.,  24.,  36.,  48.,  60.,  72.,  84.,  96., 108., 120.]])

nt.assert_array_equal(o1, multi_table(12, 10)[0])
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

<!-- grade_id: cell-aac2c2ae551e8d77, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 1 Mark
nt.assert_array_equal(4290, multi_table(12, 10)[1])
print('All Tests Passed!!!')
```

### Question 2

Write a function that takes in three scalar inputs and returns these values in a single row vector of increasing order. The function may not use any builtin functions but must perform this task only using if statements (possibly nested).

<!-- grade_id: cell-c7a5cca9a972fb86, points: 0, grade: false, solution: true -->
```python
def tri_sort(a, b, c):
    # YOUR CODE HERE
    arr=[a,b,c]
    n=3
        
    if arr[0] > arr[1]:
        arr[0], arr[1] = arr[1], arr[0]
    
    if arr[1] > arr[2]:
        arr[1], arr[2] = arr[2], arr[1]
    
    if arr[0] > arr[1]:
        arr[0], arr[1] = arr[1], arr[0]
    
                
    return arr
        
        
    raise NotImplementedError()
```

<!-- grade_id: cell-747781cad2b09589, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 1 Mark
nt.assert_array_equal(np.array([1, 5, 8]), tri_sort(5, 8, 1))
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

<!-- grade_id: cell-2fee2ba965de66e0, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# Hidden cases
# 2 Marks

```

### Question 3

Write a function that takes in one input argument, $x$. If $x$ is an empty matrix, the function must return -1. If $x$ is a scalar it must return $0$. If $x$ is a vector it must return 1. Finally, if $x$ is none of these it must return 2. You may not use any builtin function such as `ðš’ðšœðšœðšŒðšŠðš•ðšŠðš›` etc.



<!-- grade_id: cell-2153840054d0d5f3, points: 0, grade: false, solution: true -->
```python
def classifier(x):
    # YOUR CODE HERE

    
    if len(x)!=0 and type(x[0])==int:
        return 0
    
    if x.shape[0]==0:
        return -1
    
    elif x.shape[0]==1:
        return 1
    
    elif x.shape[0]>=2:
        return 2
    
    raise NotImplementedError()
```

<!-- grade_id: cell-d365fc705fd2023d, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 2 Marks
assert(2 == classifier(np.random.rand(2, 5)))
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

<!-- grade_id: cell-bf61df143c7cd376, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 1 Marks
assert(1 == classifier(np.random.rand(1, 5)))
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

<!-- grade_id: cell-4c1fe13eeb05f05d, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 2 Marks
assert(-1 == classifier(np.array([])))
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

<!-- grade_id: cell-ad56ef55fa2427fd, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 1 Marks
assert(0 == classifier([1]))
print('All Tests Passed!!!')
```

Output:
```
All Tests Passed!!!
```

## Main Exercises

Complete the following problems:

### Question 1

Write a function `fixed_steepest_descent` which implements the method of Steepest Descent using a fixed step size. The function should take as inputs the multivariate function `f`, the gradient `g`, some initial guess `x`, the fixed step size `alpha` and some tolerance `tol`. Use `norm(x, 2)` to control your tolerance. The function should return the approximation to the minimum as well as the number of iterations required to find it.


<!-- grade_id: cell-7cc3cb283ab264ac, points: 0, grade: false, solution: true -->
```python
def fixed_steepest_descent(f, g, x, alpha, tol):
    # YOUR CODE HERE
    
    it=0
    while(np.linalg.norm(x,2)>tol):
        prev_x=x
        x=x+(alpha*g(prev_x))
        it+=1
    
    return x,it
        
        
    raise NotImplementedError()
```

<!-- grade_id: cell-d39033d610ff7c90, points: 10, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 10 Marks
f1    = lambda x: 2*x[0]**2 + 3*x[1]**2
g1    = lambda x: np.array([-4*x[0], -6*x[1]])
x0    = np.array([1, 1]) 
alpha = 52/280
tol   = 1e-5

nt.assert_array_almost_equal([ 4.91552813e-06, -3.32604042e-09], fixed_steepest_descent(f1, g1, x0, alpha, tol)[0])
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

<!-- grade_id: cell-bb2827c0519a94eb, points: 5, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell 
# 5 Marks
f1    = lambda x: 2*x[0]**2 + 3*x[1]**2
g1    = lambda x: np.array([-4*x[0], -6*x[1]])
x0    = np.array([1, 1]) 
alpha = 52/280
tol   = 1e-5

assert(9 == fixed_steepest_descent(f1, g1, x0, alpha, tol)[1])
print('Test case passed!!!')
```

Output:
```
Test case passed!!!
```

```python

```